```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)

Names and Onyens of fellow students you discussed Homework4 problems and ideas:

- Students: 

We encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 4: Data Visualization, Summary Statistics, and Probability  (50 points)

Please complete this notebook by filling in the cells provided.

!! Please submit the generated **zip** file and a **PDF** of your notebook. You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop. All problems will be manually graded, so no need to submit ipynb or zip file. **But you need to make sure all your code is visible and not cut off in the pdf. If we can't see your answer code for a coding problem, points may be deducted.**

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!!


## Problem 1. Modeling (30 points)
This problem deals with determining whether a proposed model is reasonable for describing one attribute as a function of another for ROUSes.

First, you will be evaluating which of 3 models would best describe how `Length` is a function of `Age`.

Start by running the code below.

```python
import pandas as pd
import seaborn as sns
rouses_model_df = pd.read_csv("ROUSes.csv")
```

**Problem 1.1 (1 point).** Starting off by visualizing the data is always a good option.  Use seaborn `scatterplot()` to plot the $Length$ (y axis) vs. $Age$ (x axis).

```python
# write your code here, replacing ...
...
```

**Problem 1.2 (1 point).** Use seaborn `barplot()` to plot $Temperament$.

```python
# write your code here, replacing ...
...
```

**Problem 1.3 (2 point).** Use seaborn `histplot()` to plot the density of $Age$.

```python
# write your code here, replacing ...
...
```

**Problem 1.4 (2 points).**  If you were to make a density histogram with custom bins, and the first bin is [0, 5), what is the height of that bar in the histogram of $Age$? Use Python code to compute it. 

```python
# write your code here, replacing ...
...
```

**Problem 1.5 (3 points).** Now add a new feature (column) to `rouses_model_df` for each of the 3 models below, and fill in the feature with data generated by the model.  Name the features `'Length Model 1'`, `'Length Model 2'`, and `'Length Model 3'`.  Note that the first model has a square root; for this you can use `**0.5`.  Hint: if you are unsure of mathematical operation ordering, parentheses are highly recommended!

  Length Model 1 = 1.08 * $\sqrt{Age}$

  Length Model 2 = 1.0 + 0.128 * $Age$

  Length Model 3 = 1.0 + 0.00229 * $Age^2$



```python
# write your code here, replacing ...
...
```

**Problem 1.6 (4 points).** There are many ways to compare data generated from a model to the actual data.  One of the easiest is to overlay plots of the actual and generated data.  One very simple way to do this is to put multiple scatterplots in the same cell, one after the other; seaborn is smart enough to automatically change the color for each of the plots.  Do this in the below; your plot should look like this: \
 ![](LengthModel.png)

```python
# write your code here, replacing ...
...
```

**Problem 1.7 (2 points).**  Just by looking at the plots, which of the models would you say is the best fit to the actual data?  

_Type your answer here, replacing this text._

**Problem 1.8 (4 points).** Summary statistics can also be used to compare data.  For the `Length` data and the data generated by each model, use pandas to help you compute the range to the nearst tenth.  Assign the 4 range values to the corresponding variables in the cells below.

```python
length_range = ...
length_range_model1 = ...
length_range_model2 = ...
length_range_model3 = ...
```

```python
grader.check("q1.8")
```

**Problem 1.9 (4 points).** Now do the same for median.  In other words, for the `Length` data and the data generated by each model, use pandas to help you compute the median to the nearst tenth.  Assign the 4 median values to the corresponding variables in the cells below.

```python
length_median = ...
length_median_model1 = ...
length_median_model2 = ...
length_median_model3 = ...
print(length_median, length_median_model1, length_median_model2, length_median_model3)

```

```python
grader.check("q1.9")
```

**Problem 1.10 (2 points).** Considering the values you computed, which summary statistic (range or median) would be more useful in deciding the model that best approximates the data? Why? 

_Type your answer here, replacing this text._

**Problem 1.11 (2 points).** Another very useful visualization tool for comparing data is boxplot.  First, execute the code below:

```python
sns.boxplot(data=rouses_model_df)
```

```python
new_argument = ...
sns.boxplot(data=new_argument)
```

**Problem 1.12 (3 points).** Using any technique you wish, decide which of the models below is the best model for `Weight`.  Show some kind of work or visualization in the Python cell, and write what model matches best in the cell below that.

  Weight Model 1 = 51 * $\sqrt{Age}$

  Weight Model 2 = 6.9 * $Age$

  Weight Model 3 = 0.125 * $Age^2$



```python
# show your work here
...
```

_Type your answer here, replacing this text._

## Problem 2. Probability (20 points)

In this problem, you'll be analyzing the fairness of a 6-sided die.

A street vendor asks you if you'd like to play a game.  The game is:
* pay \$10 to roll a 6-sided die 3 times
* if any of those rolls is a 6 he pays you $11.  

**Problem 2.1 (2 points).** What are your chances of rolling at least one 6 in a game? (show how you calculate this)

*Write your answers here*

**Problem 2.2 (2 points).** What are your chances of rolling only one 6 in a game? (show how you calculate this)

*Write your answers here*

**Problem 2.3 (2 points).** What are your chances of winning the game given that the first two rolls are not a 6? (show how you calculate this)

*Write your answers here*

**Problem 2.4 (2 points).** What are your chances of winning the game given that the first roll is not a 6? (show how you calculate this)

*Write your answers here*

**Problem 2.5 (2 points).** What are your chances of losing the game given that the first two rolls are not a 6? (show how you calculate this)

*Write your answers here*

**Problem 2.6 (2 points).** What are your chances of losing the game given that the first roll is not a 6? (show how you calculate this)

*Write your answers here*

**Problem 2.7 (2 points).** If you play this game four times, what is the probability that you lose all four games?

*Write your answers here*

**Problem 2.8 (2 points).** What is the **expected value** of your profit or loss if you play the game once?

*Write your answers here*

**Problem 2.9 (2 points).** What is the **expected value** of your profit or loss if you play the game 10 times?

*Write your answers here*

**Problem 2.10 (2 points).** If you play the game 100 times, are you guaranteed to win at least one of the games? If not, what is the probability?

*Write your answers here*

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework4

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.8 (points: 4)
```
>>> assert length_range == 6.9
```


Tests for q1.9 (points: 4)
```
>>> assert length_median == 3.9
```
