```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4.ipynb")
```

## Name: [Your name] (your onyen)

In this activity, our focus will be on using Pandas to operate on tables. You'll find referring back to lecture 6 and 7 slides and demos helpful.

**Submission instruction**: Submit both zip file and pdf to Gradescope. Some questions will be graded by autograder, while some will be manually graded.

**Work in groups of 2~3.**

Names and onyens of people you worked with today:

- Teammate 1:
- Teammate 2:
- Teammate 3:



# Lab 4. Fun with Pandas
<img src="https://get.pxhere.com/photo/Zoo-Berlin-zoo-berlin-panda-plant-carnivore-nature-branch-biome-terrestrial-animal-natural-material-trunk-tree-snout-fur-nature-reserve-jungle-twig-tail-split-rail-fence-wildlife-grass-wood-bear-claw-1643708.jpg" width=800px>
<div style="text-align: center"> </div>

## Helpful guidelines
When working with tables on the kinds of questions in this lab and in lecture, start by thinking:
- _row selection:_ "Which rows do I need?  How do I get those rows?"
- _column selection:_ "Which columns do I need?"
- _action:_ "What do I want to do with these columns? Do I want to calculate a single number like average?  Do I need to manipulate all rows in this column to get another column of the same size?"
- _result:_ "Do I need to assign the output of this to another variable?"

Thinking through what you are trying to do step by step makes these kinds of problems much easier.

# Our solution code in parts 1-4 does NOT use an `if` statement nor a `for` loop!!

## 1. Loading a table from a file, and understanding the table. (5 minutes)
First, you will load some Comma-Separated Values (CSV) data from a file.  You'll use the pandas function `read_csv()` to do this, so pandas needs to be imported (which is done for you in the code below).  

**Q**: Add one line of code that calls `pd.read_csv()` to read a file named 'fire_swamp_flame_spurt.csv', and assign the table to a new variable `flame_df`. You can do this by replacing the ... with your code.

```python
import pandas as pd
...
flame_df
```

```python
grader.check("read")
```


![](https://blogger.googleusercontent.com/img/b/R29vZ2xl/AVvXsEiOV_tOGEuy59VbqanB3ZEvVGvst51KQxp3xi6fQmXgYprxnAcvrsW2bXgzftG8EaOpcqONeMJJG3J2ze7W7Jb834X3XcGrElKtAlu3GOxv72dFgIaetvIuvYZ78mMkxkR4PcwzwsXJanF-/s1600/The+Princess+Bride+Fire+Swamp+Fireball.jpg)

[Fire Swamp flame spurt From the movie The Princess Bride, Professor Marks' favorite movie.]

You should see table with 3 labeled columns and 272 rows containing data for the deadly Fire Swamp flame spurt.  (Think of it like a hotter version of Old Faithful).  
The Fire Swamp is said to erupt fairly regularly, and this table records the following information from the most recent 272 flame spurts.

- `duration`: Duration of the flame spurt (in minutes).
- `waiting`: Interval between flame spurts (in minutes).
- `height`: Height of the flame spurt (in feet).

**Remark**: Run the following cell to learn more about the data type of the Pandas table.

```python
type(flame_df)
```

Note that sometimes when you check the type of something, the return value looks more complex than what you would expect, but it is still equivalent to the simpler type that you were expecting, as you can see in the next cell.

```python
type(flame_df) == pd.DataFrame
```

**Q**: Run each of the following cells. Make sure to understand what kind of information they give you. Be careful - some of them are calling functions (table methods ), so require parenthesis `()`, while some are just looking up information attached to the table (think of it like metadata in your phone images), so they don't involve `()`.

```python
flame_df.dtypes
```

```python
flame_df.columns
```

```python
flame_df.shape
```

```python
flame_df.info()
```

```python
flame_df.describe()
```

## 2. Boolean indexing and adding columns (10 minutes)

**Q**: Write one line of code below that selects the column `height` in table `flame_df`. We showed you two ways to achieve this in lecture.

```python
# Write your code here!


```

It turns out that if the height of the eruption is greater than 10, there is a popping noise half a minute before the eruption happens.
Add a new column to `flame_df` named 'popping sound' that contains Boolean values for whether a popping sound would occur for each eruption.

**Q**: Write two lines of code that do the following:

1. Derive an array of boolean values for whether a popping sound would occur for each eruption. This array would have the same length as the number of rows in the `flame_df` table, and its values should be `True` or `False`. Assign this array to a new variable called `popping_sound`.

2. Add `popping_sound` as a new column in the table `flame_df`. The column name should be "popping_sound."

```python
popping_sound = ...
...
flame_df
```

```python
grader.check("popping")
```

**Q**: Use your boolean array `popping_sound` to index into only the rows of `flame_df` that has a popping sound. This is one line of code. Result should be a table with 219 rows and 4 columns. 

```python
# Write your code here!


```

**Q**: What does the tilde (`~`) do in this code?

[Write your answer here!]

```python
~popping_sound
```

```python
flame_df[~popping_sound]
```

## 3. Working with the rows of the table (5 minutes)
Often, you'll be doing operations on a single row, or a set of rows.  

**Q**: The operations in the cells below are row operations.  Check what each one does by running it, and try putting in a few different numbers.

```python
flame_df.loc[3]
```

```python
flame_df.head(4)
```

```python
flame_df.tail(5)
```

## 4. Boolean indexing, revisited (15 minutes)
Let's imagine that a robot named RobotCat is collecting this data from the Fire Swamp. RobotCat follows one algorithm - when there is a fire spurt, take measurements of the height and duration, then step away **after one hour of waiting**. Of course this means that if the interval between eruptions is shorter than one hour, RobotCat might get burned (oh no!). 

**Q**: Out of 272 flame spurts, how many times did RobotCat get burned?  Set `burned_count` to this number. 

**Hint**: My solution did this in two steps.

1. Derive an array of boolean values for whether the waiting interval is less than an hour. This array should have 272 values, which are all `True` or `False`. Assign this array to a new variable called `less_than_hour`.

2. Using the trick we saw in lecture with boolean arrays, calculate `burned_count` from the array `less_than_hour`.

```python
less_than_hour = ...
burned_count = ...
```

```python
grader.check("burned")
```

In an effort to get RobotCat burned less, you decide to equip it with an additional algorithm. Now, RobotCat will jump out of the way if either of the two things are true.

- It has waited for an hour.
- It hears the popping noise first while waiting.
  
**Q**: Now how many times does RobotCat get burned?  Assign the result to `burned_count_popping`.

Before you jump into code, first discuss with your group:

    RobotCat will get burned if less_than_hour is (True / False) (and / or) flame_df['popping sound'] is (True / False).  

Once you have the correct answer for the three choices, this simply becomes a matter of 1) combining two boolean arrays using and (`&`) or or (`|`) to create a new boolean array, 2) then using the same trick as earlier to calculate the total number of `True` in that array.

**Hint**: This new strategy should lower the number of times RobotCat gets burned! 

```python
burned_count_popping = ...
burned_count_popping
```

```python
grader.check("burned_popping")
```

## 5. Using groupby (15 minutes)
`groupby` is one of the most power pandas functions.  In order to see the full power of groupby, we will first add a new column called 'size' and initialize the entries to 'medium'.

**Q**: Run the following two cells and try to understand what's happening. No need to write any new code.


```python
flame_df['size']='Medium'
flame_df.head()
```

Next, we will set some of the values of column 'size' to either 'big' or 'small' depending on height:

```python
flame_df.loc[flame_df['height']>35,'size']='Large'
flame_df.loc[flame_df['height']<15,'size']='Small'
flame_df.head()
```

**Q**: Now you're curious what the average waiting time is for small, medium, and large eruptions. Use a `for`-loop and `groupby()` to `print` the average waiting time for each size of eruption.  

**Hint**: Here's what our solution prints out. We also rounded our values to two numbers after the decimal point for readability, which you are not required to do.

```
Large 71.84
Medium 71.19
Small 69.37
```


```python
...
```

**Q**: So based on the values, what do you think? Is the size of the eruption related to the waiting time?  When we learn more about statistics, we'll have additional tools to answer this question more quantitatively.

[Write your answer here!]

## 6. Bonus: The power of pandas
Professor Rick shows Professor Harlin the Python function below (which looks like it was written by a C programmer from the 1990s) and asks "Is there a better way to do this with Pandas? Asking for a friend." 

```python
def attribute_total( df, attribute):
    total = 0
    for attrib_item in df.columns:
        if attrib_item == attribute:
            num_rows = len(df)
            for i in range(0,num_rows):
                total = total + df.at[i, attrib_item]
    return total


```

First of all, what does it look like the function above trying to do?  Would the fuction above execute?  (a good way to answer a question like that is to try calling it!) Does the code do what it is supposed to?  Professor Harlin's answr to Professor Rick: "Your 'friend' is doing things the hard way.  That's only one line of Pandas code."  How would you do this easily with Pandas?  Write your code below, using the string `attribute` and table `df` that are defined for you.

```python
attribute = 'height'
df = flame_df
pandas_attribute_total = ...
```

```python


```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 4

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for burned (points: 2)
```
>>> assert burned_count == 77
```


Tests for burned_popping (points: 2)
```
>>> assert burned_count_popping == 17
```


Tests for popping (points: 2)
```
>>> assert flame_df.loc[0, 'popping_sound'] == True
>>> assert flame_df.loc[1, 'popping_sound'] == False
```


Tests for read (points: 1)
```
>>> assert flame_df.loc[1, 'waiting'] == 54
```
