```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

# Welcome

This is a Jupyter Notebook, which is an interactive “notebook” where we will be writing a mixture of Python code and text explanations of our results. Each of these boxes is a cell, which can contain a piece of code or text. You’ll become very familiar with it over the semester.

# 1. Hello Team

Form a group of **two to three** for this part of the activity. There's plenty of time, so please take your time and feel free to ask each other follow-up questions. 

Introduce yourselves, starting with the person who's traveled furthest away from Chapel Hill over the summer! First, share with each other the following. 

1. Your name. 
2. Your area or planned area of study at UNC. 
3. Why you are interested in Data Science.
4. What kind of group project topics you are interested in. 

Once each of you have shared your interest in Data Science move on to the following question. You might want to spend a minute or so thinking about this one.  
> *If you were a vegetable, which vegetable would you be, **and why**?* 

Make sure that you have both selected a vegetable and a reason. Note that this question is **not** the same as "what is your favorite vegetable?"

For the purposes of this question, botanical vegetables and culinary vegetables are valid choices. For example, "tomato" or "squash" are both fully acceptable responses.

# 2. Hello Jupyter

The most important thing you will need to do in Jupyter is read, write, and run Python code. Write the following code, exactly as it appears, in the *code cell* below, underneath the `# comment`. 

```python
print("to boldly go where no one has gone before")
```

Then, with your cursor still in the code cell, hit `shift + enter` on your keyboard to run the code. You may also see a "Play" or "Run" button in the Jupyter toolbar which you can use to achieve the same result and advance to the next cell. There is also a keyboard shortcut to run the code and not advance to the next cell, which is often convenient. The shortcut depends on your operating system and whether you are using Jupyter Notebook or Jupyter Lab. These shortcuts are shown in the menu bar.

```python
# write your code here

```

An important feature of the Jupyter interface is that the Python environment is persistent between code cells. That means that variables that you define in one code cell can be used in later code cells. 

In the cell below, type: 

```python
vegetables = ["corn", "carrot", "onion"]
```

Replace the vegetables with the vegetables you each chose in 1. 

Run the code cell (`shift + enter`). 

```python
# write your code here
vegetables = ...
```

```python
grader.check("q1.1")
```

The variable `vegetables` is now part of the Python session and is available for future use. For example, type the following code in the cell below and run it. 

```python
print("If we were a vegetable, we would be", vegetables, ".")
```

You'll learn more about how this works in the next week or so. 

```python
# write your code here

```

Sometimes, you'll need to write some code, but there won't be a code cell present. In this case, you can use the `+` symbol on the toolbar to add a new code cell. Add a new cell now, immediately below this text, and type the following code: 

```python
print("One of us would be a", vegetables[0], ".")
```

\[add code cell immediately below\]

On occasion, you may find it useful to *restart the kernel*. Restarting the kernel has the effect of erasing all variables in the current Python session. 

**Restart the kernel by clicking the ↻ button in the Jupyter toolbar.** If asked whether you are sure you want to restart the kernel, choose "yes" 

Create a code cell below, and type in the following code: 

```python
print("And some of us would be", vegetables[1:], ".")
```

\[add code cell immediately below\]

You should observe a `NameError: name 'vegetables' is not defined`. This is because, after restarting the kernel, the variable `vegetables` is no longer present in our namespace. To get it back, go back to the code cell in which you first defined this variable and run it. Then, the cell above will run without error. 

The toolbar also contains useful tools for duplicating cells, removing cells, and moving cells up and down. 

# 3. Hello Markdown

Throughout DATA110, you will frequently be asked to write explanatory text in which you describe your problem-solving strategy and the structure of your code. One way to do this will be in comments and docstrings, which we will introduce later. Another important way is with *Markdown cells.* Markdown (https://en.wikipedia.org/wiki/Markdown#Example) is a simple, readable way to create formatted text. It was invented by John Gruber and Aaron Swartz. This cheatsheet (https://commonmark.org/help/) contains a useful and concise set of instructions for Markdown editing. 

To create a Markdown cell, first create a code cell like you did before. Then, with your cursor in the code cell, click the dropdown menu in the toolbar that says "Code".  Select "Markdown" from this menu instead. 

Return to the cell and type the following characters, exactly as they appear:

`*To boldly go* **where no one has gone before.**`

Then, "run" the Markdown cell (`shift + enter` again) to see your text rendered with italic and bold formatting. 

Now, work with your group to reproduce in Markdown the following algorithm for enjoying a relaxing cup of tea. 

<br>

<figure class="image" style="width:40%">
  <img src="https://raw.githubusercontent.com/PhilChodrow/PIC16A/master/_images/making-tea.png" alt="A list of six items describing how to make tea. The list reads:             
    1. Boil water.
    2. Place tea bag in mug.
    3. Pour water in mug. **Be careful, water is hot!**
    3. *Optional: add milk and/or sugar*. 
    4. Drink tea. 
        - `while cup is not empty`
            - Carefully raise mug to mouth. 
            - Take a sip. 
    5. Clean mug. ">
  <figcaption><i></i></figcaption>
</figure>


\[add markdown cell immediately below\]

# 4. Goodbye first recitation

Great job! It's time to submit the worksheet. Make sure to run all cells (you can click on the fastforward button above (>>)), download the zip file, and submit it on Gradescope under `Lab 1`.  


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Gradescope Lab1

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1 (points: 10)
```
>>> assert type(vegetables) == list
```
